---
title       : STAT5
description : 
attachments :



--- type:NormalExercise lang:r xp:100 skills:1 key:509c602d6a
## One sample t test

** Are our students taller than average? **

The *one sample t test* compares the mean of one sample to a fixed value.

For example, the WHO 2007 reference height for 19 year-old females has a mean of 163.1548 cm. 



*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
